// intro2.rs
//
// Make the code print a greeting to the world.
//
// Execute `rustlings hint intro2` or use the `hint` watch subcommand for a hint.
// Execute `rustlings note intro2` or use the `note intro2` watch subcommand for lzz's note.



// I AM NOT DONE

fn main() {
    println!("Hello {}!");
}
