// variables5.rs
//
// Execute `rustlings hint variables5` or use the `hint` watch subcommand for a hint.
// Execute `rustlings note variables5` or use the `note variables5` watch subcommand for lzz's note.



// I AM NOT DONE

fn main() {
    let number = "T-H-R-E-E"; // don't change this line
    println!("Spell a Number : {}", number);
    number = 3; // don't rename this variable
    println!("Number plus two is : {}", number + 2);
}
