// variables1.rs
//
// Make me compile!
//
// Execute `rustlings hint variables1` or use the `hint` watch subcommand for a hint.
// Execute `rustlings note variables1` or use the `note variables1` watch subcommand for lzz's note.



// I AM NOT DONE

fn main() {
    x = 5;
    println!("x has the value {}", x);
}
