// variables2.rs
//
// Execute `rustlings hint variables2` or use the `hint` watch subcommand for a hint.
// Execute `rustlings note variables2` or use the `note variables2` watch subcommand for lzz's note.



// I AM NOT DONE

fn main() {
    let x;
    if x == 10 {
        println!("x is ten!");
    } else {
        println!("x is not ten!");
    }
}
