// variables4.rs
//
// Execute `rustlings hint variables4` or use the `hint` watch subcommand for a hint.
// Execute `rustlings note variables4` or use the `note variables4` watch subcommand for lzz's note.



// I AM NOT DONE

fn main() {
    let x = 3;
    println!("Number {}", x);
    x = 5; // don't change this line
    println!("Number {}", x);
}
