// macros2.rs
//
// Execute `rustlings hint macros2` or use the `hint` watch subcommand for a hint.
// Execute `rustlings note macros2` or use the `note macros2` watch subcommand for lzz's note.



// I AM NOT DONE

fn main() {
    my_macro!();
}

macro_rules! my_macro {
    () => {
        println!("Check out my macro!");
    };
}
