// macros1.rs
//
// Execute `rustlings hint macros1` or use the `hint` watch subcommand for a hint.
// Execute `rustlings note macros1` or use the `note macros1` watch subcommand for lzz's note.



// I AM NOT DONE

macro_rules! my_macro {
    () => {
        println!("Check out my macro!");
    };
}

fn main() {
    my_macro();
}
