// functions2.rs
//
// Execute `rustlings hint functions2` or use the `hint` watch subcommand for a hint.
// Execute `rustlings note functions2` or use the `note functions2` watch subcommand for lzz's note.



// I AM NOT DONE

fn main() {
    call_me(3);
}

fn call_me(num:) {
    for i in 0..num {
        println!("Ring! Call number {}", i + 1);
    }
}
